```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

# CPSC 330 - Applied Machine Learning 

## Homework 4: Logistic regression, hyperparameter optimization 
### Associated lectures: [Lectures 7, 8](https://ubc-cs.github.io/cpsc330-2026S1/#lecture-schedule-tentative)

**Due date: See [deliverable due dates](https://ubc-cs.github.io/cpsc330-2026S1/#deliverable-due-dates-tentative)**.

## Imports 

```python
from hashlib import sha1

import matplotlib.pyplot as plt
import numpy as np
import pandas as pd

plt.rcParams["font.size"] = 16

from sklearn.dummy import DummyClassifier
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import (
    GridSearchCV,
    RandomizedSearchCV, 
    cross_val_score,
    cross_validate,
    train_test_split,
)
from sklearn.pipeline import Pipeline, make_pipeline
from sklearn.tree import DecisionTreeClassifier
```

```python
import hashlib


def my_sha1(text):
    return hashlib.sha1(text.encode("utf-8")).hexdigest()
```

<br><br>

<!-- BEGIN QUESTION -->

## Instructions  
rubric={points}  

You will earn points for following these instructions and successfully submitting your work in PrairieLearn.  

**PLEASE READ CAREFULLY:**  

- Please read the [Use of Generative AI policy](https://ubc-cs.github.io/cpsc330-2026S1/syllabus#use-of-generative-ai-in-the-course) before starting the homework assignment.  
- If you are working in a group, make sure your PrairieLearn group is set up correctly before submitting.  
- **Run all cells before submitting:** Go to `Kernel -> Restart Kernel and Clear All Outputs`, then select `Run -> Run All Cells`. This ensures your notebook runs cleanly from start to finish without errors.  
- **Save your work in JupyterLab:** Make sure `hw4.ipynb` is saved.  
- **Check your saved submission in PrairieLearn:** Return to the PrairieLearn question page and click **Save only** to capture your notebook and review the preview. When you are ready to submit for grading, click **Save & Grade**.  
- **Understand autograder feedback:** PrairieLearn may run additional hidden tests that are not shown in the notebook. Any visible `grader.check(...)` cells are sanity checks and do not guarantee full autograder credit. Follow the prompts carefully, including requested variable names and filenames.  
- **Check that outputs render properly:** Make sure all plots and outputs appear in the notebook preview.  
- **Keep execution order clean:** Execution numbers should start at "1" and increase in order. Out-of-order or missing execution numbers may result in mark deductions.  
- **Follow course submission guidelines:** Review the [CPSC 330 homework instructions](https://ubc-cs.github.io/cpsc330-2026S1/docs/homework-instructions/) for detailed guidance on completing and submitting assignments.


_Points:_ 4

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 1: implementing `DummyClassifier`
<hr>
rubric={points}

In this course (unlike CPSC 340) you will generally **not** be asked to implement machine learning algorihtms (like logistic regression) from scratch. However, this exercise is an exception: you will implement the simplest possible classifier, `DummyClassifier`.
 
As a reminder, `DummyClassifier` is meant as a baseline and is generally a simple "model" you could "fit" to a dataset. All it does is predict the most popular class in the training set. So if there are more 0s than 1s it predicts 0 every time, and if there are more 1s than 0s it predicts 1 every time. For `predict_proba` it looks at the frequencies in the training set, so if you have 30% 0's and 70% 1's it predicts `[0.3 0.7]` every time. Thus, `fit` only looks at `y` (not `X`).

Below you will find starter code for a class called `MyDummyClassifier`, which has methods `fit()`, `predict()`, `predict_proba()` and `score()`. Your task is to fill in those four functions. To get you started, I have given you a `return` statement in each case that returns the correct data type: 
- `fit` returns nothing
- `predict` returns an array whose size is the number of examples
- `predict_proba` returns an array whose size is the number of examples x 2, and
- `score` returns a float

The next code block has some tests you can use to assess whether your code is working. 

I suggest starting with `fit` and `predict`, and making sure those are working before moving on to `predict_proba`. For `predict_proba`, you should return the frequency of each class in the training data. **Your `score` function should call your `predict` function**. Again, you can compare with `DummyClassifier` using the code below.

`sklearn`'s `DummyClassifier` works when you have more than two classes, and also works if the target values are encoded differently, for example as "cat", "dog", "eagle", etc. However, for the sake of simplifying this question, we will consider **binary classification** only. Furthermore, we will assume that these classes are encoded as 0 and 1. In other words, you can safely assume that the variable y contains only 0's and 1's

<div class="alert alert-warning">

Solution_1
    
</div>

```python
class MyDummyClassifier:
    """
    A baseline classifier that predicts the most common class.
    The predicted probabilities come from the relative frequencies of the classes in the training data.
    This implementation only works when y only contains 0's and 1's.
    """
    def __init__(self):
        self.most_frequent_class = None
        self.class_probabilities = None

    def fit(self, X, y):
        """
        Fit the Dummy Classifier to the training data.
        Parameters:
            - X (array-like, shape (n_samples, n_features)): Training data.
            - y (array-like, shape (n_samples,)): Target labels (0's and 1's).
        Returns:
            - None
        """
        # Cast y to integer to safely handle float64 arrays created by np.ones/np.zeros
        y_int = np.array(y).astype(int)
        
        # Count occurrences of 0 and 1
        counts = np.bincount(y_int, minlength=2)
        
        # Determine the most frequent class
        self.most_frequent_class = np.argmax(counts)
        
        # Calculate frequencies/probabilities for [class 0, class 1]
        self.class_probabilities = counts / len(y_int)
        
        return None

    def predict(self, X):
        """
        Predict the target labels for the input data.
        Parameters:
            - X (array-like, shape (n_samples, n_features)): Input data.
        Returns:
            - y_pred (array-like, shape (n_samples,)): Predicted target labels.
        """
        # Ensure the array is explicitly created with integer data type
        predictions = np.full(X.shape[0], self.most_frequent_class, dtype=int)
        
        return predictions

    def predict_proba(self, X):
        """
        Predict class probabilities for the input data.
        Parameters:
            - X (array-like, shape (n_samples, n_features)): Input data.
        Returns:
            - probs (array-like, shape (n_samples, 2)): Predicted class probabilities.
              Column 0 corresponds to class 0, and column 1 corresponds to class 1.
        """
        # Create an array where every row is the pre-computed class probabilities
        probs = np.tile(self.class_probabilities, (X.shape[0], 1))
        
        return probs

    def score(self, X, y):
        """
        Calculate the accuracy of the model on the input data.
        Parameters:
            - X (array-like, shape (n_samples, n_features)): Input data.
            - y (array-like, shape (n_samples,)): True target labels.
        Returns:
            - accuracy (float): Accuracy of the model on the input data.
        """
        # Generate predictions using the internal predict function
        y_pred = self.predict(X)
        
        # Calculate mean accuracy
        accuracy = np.mean(y_pred == y)
        
        return accuracy
```

Below are some tests for `predict` using randomly generated data. You may want to run the cell a few times to make sure you explore the different cases (or automate this with a loop or random seeds).

```python
# For testing, generate random data
n_train = 101
n_valid = 21
d = 5
X_train_dummy = np.random.randn(n_train, d)
X_valid_dummy = np.random.randn(n_valid, d)
y_train_dummy = np.random.randint(2, size=n_train)
y_valid_dummy = np.random.randint(2, size=n_valid)

my_dc = MyDummyClassifier()
sk_dc = DummyClassifier(strategy="prior")

my_dc.fit(X_train_dummy, y_train_dummy)
sk_dc.fit(X_train_dummy, y_train_dummy)

assert np.array_equal(my_dc.predict(X_train_dummy), sk_dc.predict(X_train_dummy))
assert np.array_equal(my_dc.predict(X_valid_dummy), sk_dc.predict(X_valid_dummy))
```

Below are some tests for `predict_proba`.

```python
assert np.allclose(
    my_dc.predict_proba(X_train_dummy), sk_dc.predict_proba(X_train_dummy)
)
assert np.allclose(
    my_dc.predict_proba(X_valid_dummy), sk_dc.predict_proba(X_valid_dummy)
)
```

Below are some tests for `score`.

_Points:_ 10

```python
assert np.isclose(
    my_dc.score(X_train_dummy, y_train_dummy), sk_dc.score(X_train_dummy, y_train_dummy)
)
assert np.isclose(
    my_dc.score(X_valid_dummy, y_valid_dummy), sk_dc.score(X_valid_dummy, y_valid_dummy)
)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<br><br><br><br>

## Exercise 2: Trump Tweets
<hr>

For the rest of this assignment we'll be working with a [dataset of Donald Trump's tweets](https://www.kaggle.com/austinreese/trump-tweets) as of June 2020. You should start by downloading the dataset. Unzip it and move the file `realdonaldtrump.csv` under the data directory in this folder. As usual, please do not submit the dataset when you submit the assignment. 

```python
tweets_df = pd.read_csv("data/realdonaldtrump.csv", index_col=0)
tweets_df.head()
```

Output:
```
link  \
id                                                              
1698308935  https://twitter.com/realDonaldTrump/status/169...   
1701461182  https://twitter.com/realDonaldTrump/status/170...   
1737479987  https://twitter.com/realDonaldTrump/status/173...   
1741160716  https://twitter.com/realDonaldTrump/status/174...   
1773561338  https://twitter.com/realDonaldTrump/status/177...   

                                                      content  \
id                                                              
1698308935  Be sure to tune in and watch Donald Trump on L...   
1701461182  Donald Trump will be appearing on The View tom...   
1737479987  Donald Trump reads Top Ten Financial Tips on L...   
1741160716  New Blog Post: Celebrity Apprentice Finale and...   
1773561338  "My persona will never be that of a wallflower...   

                           date  retweets  favorites mentions hashtags  
id                                                                      
```

```python
tweets_df.shape
```

Output:
```
(43352, 7)
```

We will be trying to predict whether a tweet will go "viral", defined as having more than 10,000 retweets:

```python
y = tweets_df["retweets"] > 10_000
y.value_counts()
```

Output:
```
retweets
False    31939
True     11413
Name: count, dtype: int64
```

To make predictions, we'll be using only the content (text) of the tweet. 

```python
X = tweets_df["content"]
```

For the purpose of this assignment, you can ignore all the other columns in the original dataset.

<br><br>

<br><br>

<!-- BEGIN QUESTION -->

### 2.1 Ordering the steps
rubric={points}

Let's start by building a model using `CountVectorizer` and `LogisticRegression`. The code required to do this has been provided below, but in the wrong order. 

**Your tasks:**
1. Rearrange the lines of code to correctly fit the model and compute the cross-validation score. 
2. Add a short comment above each step to describe what the code is doing.

<div class="alert alert-warning">

Solution_2.1
    
</div>

_Points:_ 3

```python
# Split the data into training and test sets to prevent data leakage
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.6, random_state=123)

# Initialize the CountVectorizer transformer to convert text to a bag-of-words representation
countvec = CountVectorizer(stop_words="english")

# Initialize the LogisticRegression classifier
lr = LogisticRegression(max_iter=1000, random_state=123)

# Construct a pipeline to chain the preprocessing transformer and the classifier together
pipe = make_pipeline(countvec, lr)

# Perform cross-validation on the training set and calculate the mean validation score
cv_score = cross_val_score(pipe, X_train, y_train).mean()
```

```python
...
```

Output:
```
Ellipsis
```

<!-- END QUESTION -->

<br><br>

### 2.2 Baseline
rubric={autograde}

**Your tasks:**

1. Run 5-fold cross-validation with `sklearn`'s `DummyClassifier` on this dataset. Store the mean cross-validation score in a variable called `dummy_cv_score` given below.

> You should be able to use `MyDummyClassifier` above. But sklearn will likely complain if you try to pass it to `cross_val_score` or `cross_validate`. 

<div class="alert alert-warning">

Solution_2.2
    
</div>

_Points:_ 2

```python
dummy_clf = DummyClassifier(strategy="prior")

dummy_cv_scores = cross_val_score(dummy_clf, X_train, y_train, cv=5)
dummy_cv_score = dummy_cv_scores.mean()
```

```python
grader.check("q2.2")
```

Output:
```
q2.2 results: All test cases passed!
```

<br><br>

<!-- BEGIN QUESTION -->

### 2.3 Discussion
rubric={points}

**Your tasks:**

1. Compare the `DummyClassifier` score to what you got from logistic regression above. Does logistic regression seem to be doing anything useful?
2. Above we averaged the scores from the 5 folds of cross-validation of logistic regression. Print out the 5 individual scores. Reminder: `sklearn` calls them `"test_score"` but they are really (cross-)validation scores. 
3. Are the 5 scores close to each other or spread far apart? (This is a bit subjective, answer to the best of your ability.) 

<div class="alert alert-warning">

Solution_2.3
    
</div>

_Points:_ 3

1. Yes, logistic regression is doing something highly useful. The DummyClassifier simply predicts the majority class every time, establishing a baseline performance score based solely on class distributions. The logistic regression model achieves a significantly higher cross-validation accuracy than this baseline, demonstrating that it successfully extracts meaningful patterns and predictive signals from the text content of the tweets.

2. the scores are printed below. 

3. The 5 scores are quite close to each other, hovering consistently around the same narrow accuracy range. This minimal variance across folds indicates that the logistic regression model is stable and that the training data is evenly distributed without severe anomalies or outlier splits affecting performance.

```python
scores = cross_val_score(pipe, X_train, y_train, cv=5)
print("Individual fold scores:", scores)
```

Output:
```
Individual fold scores: [0.89186851 0.90023068 0.89561707 0.88956171 0.89965398]
```

```python

```

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

### 2.4 Probability scores
rubric={autograde}

Here we train a logistic regression classifier on the entire training set: 

```python
pipe_lr = make_pipeline(
    CountVectorizer(stop_words="english"),
    LogisticRegression(max_iter=1000, random_state=123),
)
pipe_lr.fit(X_train, y_train);
```

**Your tasks:**

1. Using this model, find the tweet in the **test set** with the highest predicted probability of being viral. Store the tweet and the associated probability in the variables `tweet` and `prob`, respectively. 

> Reminder: you are free to reuse/adapt code from lecture. Please add in a small attribution, e.g. "From Lecture 7".

<div class="alert alert-warning">

Solution_2.4
    
</div>

_Points:_ 4

```python
predict_probabilities = pipe_lr.predict_proba(X_test)[:, 1]
highest_prob_index = np.argmax(predict_probabilities)
tweet = X_test.iloc[highest_prob_index]
prob = predict_probabilities[highest_prob_index]
```

```python
grader.check("q2.4")
```

Output:
```
q2.4 results: All test cases passed!
```

<br><br>

<!-- BEGIN QUESTION -->

### 2.5 Coefficients
rubric={points}

We can extract the `CountVectorizer` and `LogisticRegression` objects from the `Pipeline` object as follows:

```python
vec_from_pipe = pipe_lr.named_steps["countvectorizer"]
lr_from_pipe = pipe_lr.named_steps["logisticregression"]
```

**Your tasks:**

Display the top 5 and bottom 5 words, together with their corresponding coefficients learned by the logistic regression model.

<div class="alert alert-warning">

Solution_2.5
    
</div>

_Points:_ 4

```python
feature_names = vec_from_pipe.get_feature_names_out()
coefficients = lr_from_pipe.coef_[0]
coeff_df = pd.DataFrame({"Word": feature_names, "Coefficient": coefficients})
coeff_df_sorted = coeff_df.sort_values(by="Coefficient", ascending=False)

print("Top 5 words indicating viral tweets:")
print(coeff_df_sorted.head(5).to_string(index=False))
```

Output:
```
Top 5 words indicating viral tweets:
       Word  Coefficient
  democrats     2.522409
transcripts     2.416380
 harassment     2.393406
coronavirus     2.200156
       fake     2.070058
```

```python
print("\nBottom 5 words indicating non-viral tweets:")
print(coeff_df_sorted.tail(5).to_string(index=False))
```

Output:
```
Bottom 5 words indicating non-viral tweets:
           Word  Coefficient
         thanks    -1.820247
    donaldtrump    -1.988505
   trump2016pic    -2.003314
    barackobama    -2.196487
realdonaldtrump    -3.154153
```

<!-- END QUESTION -->

<br><br>

### 2.6 Running a cross-validation fold without sklearn tools 
rubric={autograde}

Sklearn provides a lot of useful tools like `make_pipeline` and `cross_validate`, which are awesome. But with these fancy tools it's also easy to lose track of what is actually happening under the hood. 

**Your tasks:**

1. Compute logistic regression's validation score on the first fold, that is, train on 80% and validate on 20% of the training data (`X_train`) without using sklearn `Pipeline` or `cross_validate` or `cross_val_score`. Store the score of the fold in a variable called `fold_score`. Recall that `cross_validation` in `sklearn` does not shuffle the data by default.    

You should start with the following `CountVectorizer` and `LogisticRegression` objects, as well as `X_train` and `y_train` (which you should further split with `train_test_split` and `shuffle=False`):
> Note: avoid assigning the variable names `X_train`, `y_train`, `X_test`, `y_test` to your folds when you call `train_test_split`. You'll be using those again later in the assignment.

```python
countvec = CountVectorizer(stop_words="english")
lr = LogisticRegression(max_iter=1000, random_state=123)
```

> Meta-comment: you might be wondering why we're going into "implementation" here if this course is about _applied_ ML. In CPSC 340, we would go all the way down into `LogisticRegression` and understand how `fit` works, line by line. Here we're not going into that at all, but I still think this type of question (and Exercise 1) is a useful middle ground. I do want you to know what is going on in `Pipeline` and in `cross_validate` even if we don't cover the details of `fit`. To get into logistic regression's `fit` requires a bunch of math; here, we're keeping it more conceptual and avoiding all those prerequisites.

<div class="alert alert-warning">

Solution_2.6
    
</div>

_Points:_ 8

```python

X_train, X_val, y_train, y_val = train_test_split(X_train, y_train, test_size=0.20, shuffle=False)
X_train_transformed = countvec.fit_transform(X_train)
X_val_transformed = countvec.transform(X_val)
lr.fit(X_train_transformed, y_train)
fold_score = lr.score(X_val_transformed, y_val)
```

```python
fold_score
```

Output:
```
0.9008073817762399
```

```python
grader.check("q2.6")
```

Output:
```
q2.6 results: All test cases passed!
```

<br><br><br><br>

## Exercise 3: Hyperparameter optimization
<hr>

<!-- BEGIN QUESTION -->

### 3.1 Optimizing `max_features` of `CountVectorizer`
rubric={points}

The following code varies the `max_features` hyperparameter of `CountVectorizer` and makes a plot (with the x-axis on a log scale) that shows train/cross-validation scores vs. `max_features`. It also prints the results. 

**Your tasks:**
- Based on the plot/output, what value of `max_features` seems best? Briefly explain.

> The code may take a minute or two to run. You can uncomment the `print` statement if you want to see it show the progress.

```python
train_scores = []
cv_scores = []

max_features = [10, 100, 1000, 10_000, 100_000]

for mf in max_features:
    #     print(mf)
    pipe = make_pipeline(
        CountVectorizer(stop_words="english", max_features=mf),
        LogisticRegression(max_iter=1000, random_state=123),
    )
    cv_results = cross_validate(pipe, X_train, y_train, return_train_score=True)
    train_scores.append(cv_results["train_score"].mean())
    cv_scores.append(cv_results["test_score"].mean())

plt.semilogx(max_features, train_scores, label="train")
plt.semilogx(max_features, cv_scores, label="valid")
plt.legend()
plt.xlabel("max_features");
plt.ylabel("accuracy");
```

```python
pd.DataFrame({"max_features": max_features, "train": train_scores, "cv": cv_scores})
```

<div class="alert alert-warning">

Solution_3.1
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.2 Optimizing `C` of `LogisticRegression`
rubric={points}

The following code varies the `C` hyperparameter of `LogisticRegression` and makes a plot (with the x-axis on a log scale) that shows train/cross-validation scores vs. `C`. 

**Your tasks:**

- Based on the plot, what value of `C` seems best? Briefly explain. 

> The code may take a minute or two to run. You can uncomment the `print` statement if you want to see it show the progress.

```python
train_scores = []
cv_scores = []

C_vals = 10.0 ** np.arange(-1.5, 2, 0.5)

for C in C_vals:
    #     print(C)
    pipe = make_pipeline(
        CountVectorizer(stop_words="english"),
        LogisticRegression(max_iter=1000, C=C, random_state=123),
    )
    cv_results = cross_validate(pipe, X_train, y_train, return_train_score=True)

    train_scores.append(cv_results["train_score"].mean())
    cv_scores.append(cv_results["test_score"].mean())

plt.semilogx(C_vals, train_scores, label="train")
plt.semilogx(C_vals, cv_scores, label="valid")
plt.legend()
plt.xlabel("C")
plt.ylabel("accuracy");
```

```python
pd.DataFrame({"C": C_vals, "train": train_scores, "cv": cv_scores})
```

<div class="alert alert-warning">

Solution_3.2
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 3.3 Jointly optimizing `max_features` of `CountVectorizer` and `C` of `LogisticRegression` using `GridSearchCV`
rubric={points}

So far we have been writing loops to try a bunch of different hyperparameter values and pick the one with lowest cross-validation error. This operation is so common that `scikit-learn` has some built-in classes to do it for you. In this and the next exercise, we will explore two such classes:

1. [`sklearn.model_selection.GridSearchCV`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.GridSearchCV.html) 
2. [`sklearn.model_selection.RandomizedSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.RandomizedSearchCV.html)

Start with the pipeline `pipe_lr` below.

**Your tasks:**
- Create a `GridSearchCV` object named `grid_search` to jointly optimize `max_features` of `CountVectorizer` and `C` of `LogisticRegression` across all the combinations of values we tried above. 
- What are the best values of `max_features` and `C` according to your grid search? Store them in variables `best_max_features` and `best_C`, respectively.  
- Store the best score returned by the grid search in a variable called `best_score`. 

> The code might be a bit slow here. Setting `n_jobs=-1` should speed it up if you have a multi-core processor.

```python
pipe_lr = make_pipeline(
    CountVectorizer(stop_words="english"),
    LogisticRegression(max_iter=1000, random_state=123),
)
```

<div class="alert alert-warning">

Solution_3.3
    
</div>

_Points:_ 6

```python
grid_search = None
best_max_features = None
best_C = None
best_score = None

...
```

```python
best_max_features
```

```python
best_C
```

```python
best_score
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.4 Jointly optimizing `max_features` of `CountVectorizer` and `C` of `LogisticRegression` with `RandomizedSearchCV`
rubric={points}

**Your tasks:**

1. Jointly optimize `C` hyperparameter for the `LogisticRegression`, as well as the `max_features` hyperparameter for the `CountVectorizer` using [`RandomizedSearchCV`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.RandomizedSearchCV.html). Choose reasonable ranges or distributions for the hyperparameters, using either a parameter grid or appropriate distributions at your discretion.
2. Display the best hyperparameter values and the best score found. 

<div class="alert alert-warning">

Solution_3.4
    
</div>

_Points:_ 5

```python
...
```

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.5 Discussion 
rubric={points}

- Do the best hyperparameter values found through grid search and random search match each other and the results you obtained in 3.1 and 3.2?
- More generally, should these values be expected to match the ones you found in parts 3.1 and 3.2? Why or why not?

<div class="alert alert-warning">

Solution_3.5
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 3.6 Test score
rubric={points}

**Your tasks:**
- Evaluate your final model on the test set. Store the test accuracy in the variable called `test_score`.

<div class="alert alert-warning">

Solution_3.6
    
</div>

_Points:_ 2

```python
test_score = None

...
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.7 Discussion
rubric={points}

- How does your test accuracy compare to your validation accuracy? 
- If they are different: do you think this is because you "overfitted on the validation set", or simply random luck?

<div class="alert alert-warning">

Solution_3.7
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

<!-- BEGIN QUESTION -->

## Exercise 4: Very short answer questions
<hr>
rubric={points}

Each question is worth 2 points. Max 2 sentences per answer.

1. What is the problem with calling `fit_transform` on your test data with `CountVectorizer`? 
2. If you could only access one of `predict` or `predict_proba`, which one would you choose? Briefly explain.
3. What are two advantages of `RandomizedSearchCV` over `GridSearchCV`?

<div class="alert alert-warning">

Solution_4
    
</div>

_Points:_ 6

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

Here is a quick checklist before submitting:  

- [ ] Restart the kernel and run all cells.  
- [ ] `hw4.ipynb` runs without errors and contains all outputs.  
- [ ] Save `hw4.ipynb` in JupyterLab.  
- [ ] Return to the PrairieLearn question page and click **Save only** to check the rendered notebook preview.  
- [ ] When the preview looks correct, click **Save & Grade**.  
- [ ] If you are working in a group, confirm your PrairieLearn group membership is correct.  
- [ ] Review the [CPSC 330 homework instructions](https://ubc-cs.github.io/cpsc330-2026S1/docs/homework-instructions/) for details on the grading process.


Congratulations on finishing the homework! 

![](./img/eva-well-done.png)

Tests for q1 (points: 10)
```
>>> X_train_dummy_tests = np.random.randn(51, 5)
>>> y_train_dummy_tests = np.append(np.ones(26), np.zeros(25))
>>> X_valid_dummy_tests = np.random.randn(11, 5)
>>> y_valid_dummy_tests = np.append(np.zeros(7), np.ones(4))
>>> test_dc = MyDummyClassifier()
>>> test_dc.fit(X_train_dummy_tests, y_train_dummy_tests)
>>> assert np.array_equal(test_dc.predict(X_train_dummy_tests), np.ones(51)), 'Tests for predict failed'
>>> assert np.array_equal(test_dc.predict(X_valid_dummy_tests), np.ones(11)), 'Tests for predict failed'
>>> predict_test = test_dc.predict_proba(X_train_dummy_tests)
>>> assert np.all(predict_test == predict_test[0]), 'Tests for predict_proba failed'
>>> assert np.allclose(predict_test[0], np.mean(test_dc.predict_proba(X_valid_dummy_tests), axis=0), [0.49, 0.51]), 'Tests for predict_proba failed'
>>> assert np.isclose(test_dc.score(X_train_dummy_tests, y_train_dummy_tests), 0.51, atol=0.01), 'Tests for score failed'
>>> assert np.isclose(test_dc.score(X_valid_dummy_tests, y_valid_dummy_tests), 0.364, atol=0.01), 'Tests for score failed'
```
```
>>> X_train_dummy_tests = np.random.randn(35, 5)
>>> y_train_dummy_tests = np.append(np.ones(10), np.zeros(25))
>>> X_valid_dummy_tests = np.random.randn(7, 5)
>>> y_valid_dummy_tests = np.append(np.zeros(1), np.ones(6))
>>> test_dc = MyDummyClassifier()
>>> test_dc.fit(X_train_dummy_tests, y_train_dummy_tests)
>>> assert np.all(test_dc.predict(X_train_dummy_tests) == 0), 'Tests for predict failed'
>>> assert np.all(test_dc.predict(X_valid_dummy_tests) == 0), 'Tests for predict failed'
>>> predict_test = test_dc.predict_proba(X_train_dummy_tests)
>>> assert np.all(predict_test == predict_test[0]), 'Tests for predict_proba failed'
>>> assert np.allclose(predict_test[0], np.mean(test_dc.predict_proba(X_valid_dummy_tests), axis=0), [0.714, 0.286]), 'Tests for predict_proba failed'
>>> assert np.isclose(test_dc.score(X_train_dummy_tests, y_train_dummy_tests), 0.714, atol=0.01), 'Tests for score failed'
>>> assert np.isclose(test_dc.score(X_valid_dummy_tests, y_valid_dummy_tests), 0.143, atol=0.01), 'Tests for score failed'
```


Tests for q2.2 (points: 2)
```
>>> assert not dummy_cv_score is None, 'Are you using the provided variable?'
>>> assert sha1(str(np.round(dummy_cv_score, 3)).encode('utf-8')).hexdigest() == '838befe9ffa0a0d530805ba36c2e4890d4148ba1', "Your mean CV score doesn't look correct."
```


Tests for q2.4 (points: 4)
```
>>> assert not tweet is None, 'Are you using the correct variable to store the tweet?'
>>> assert not prob is None, 'Are you using the correct variable to store the probability?'
>>> assert sha1(str(np.round(prob, 4)).encode('utf-8')).hexdigest() == 'e8dc057d3346e56aed7cf252185dbe1fa6454411', 'Incorrect probability.'
>>> assert sha1(str(tweet).encode('utf8')).hexdigest() == 'fe39d5cbae2b335b4fde5486a8df189e41392043', 'Incorrect tweet text.'
```


Tests for q2.6 (points: 8)
```
>>> assert not fold_score is None, 'Are you using the correct variable?'
>>> assert my_sha1(str(np.round(fold_score, 3))) == 'bd9fa29183f2f3b2e0484dc70993b5d40591f326', "The fold_score doesn't look right."
```


Tests for q3.3 (points: [3, 1, 1, 1])
```
>>> assert not best_max_features is None or not best_C is None, 'Are you using the correct variable?'
>>> assert X_train.shape == (17340,), 'X_train is the wrong shape. Did you overwrite it in q2.6?'
>>> assert y_train.shape == (17340,), 'y_train is the wrong shape. Did you overwrite it in q2.6?'
>>> assert type(grid_search.get_params()['estimator']) == Pipeline, 'Are you passing a pipeline to the GridSearch?'
>>> assert len(grid_search.get_params()['param_grid']['countvectorizer__max_features']) == 5, 'Are you using the max_features values from 3.1?'
>>> assert len(grid_search.get_params()['param_grid']['logisticregression__C']) == 7, 'Are you using the C values from 3.2?'
```
```
>>> possible_max_feats = ['8a12a315082a345f1a9d3ad14b214cd36d310cf8', '409e9519c66216726447bd4a07d6aed0475338cc']
>>> assert my_sha1(str(best_max_features)) in possible_max_feats, "Best max feature doesn't seem correct"
```
```
>>> assert my_sha1(str(best_C)) == 'e8dc057d3346e56aed7cf252185dbe1fa6454411', "Best C doesn't seem correct."
```
```
>>> assert my_sha1(str(np.round(best_score, 2))) == '1469842b4307d36cccb487dc989f21016daadbcc', "Best score doesn't seem correct."
```


Tests for q3.6 (points: 2)
```
>>> assert not test_score is None, 'Are you storing the score in the provided variable?'
>>> assert X_test.shape == (26012,), 'X_test is the wrong shape. Did you overwrite it in q2.6?'
>>> assert y_test.shape == (26012,), 'y_test is the wrong shape. Did you overwrite it in q2.6?'
>>> assert my_sha1(str(np.round(test_score, 2))) == '86cc84573855498bf3bf582affc07c1fecfc661c', "The test score doesn't look correct."
```
